# Analiza danych przestrzennych - wiczenia laboratoryjne 2023/2024

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania i oceny notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Je偶eli zestaw zada wymaga skorzystania z fragment贸w kodu opracowanego w ramach wczeniejszych zestaw贸w zada nale偶y je umieci we wskazanej kom贸rce notatnika.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu, odpowiedzi uzupeniane manualnie nie podlegaj ocenie.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj na celu wyapa podstawowe bdy. Przejcie przez kod wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie i zostanie ocenione na maksymaln liczb punkt贸w.
- Zadania nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zadania, kt贸re powoduj wywietlenie komunikatu o bdzie przerywajcym wykonywanie kodu nie podlegaj ocenie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem). Rozmiar poziomy figur nie powinien przekracza 20 cali.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (nazwa zespou oraz imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespou:
Czonkowie:

---

## Zestaw zada 4: Wybrane teoretyczne rozkady prawdopodobiestwa oraz elementy teorii estymacji (Notatnik 3/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakiet贸w wchodzcych w skad standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakiet贸w
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)
```

### Dane do zada

W celu wygenerowania danych wykorzystywanych w zawartych w notatniku zadaniach i kom贸rkach testowych wykonaj poni偶sze kom贸rki.

```python
# Dane do test贸w 1
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
test_data_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
# Dane do test贸w 2
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=25, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=25, random_state=13)
test_data_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["x", "y"])
test_data_2.head()
```

Output:
```
x         y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

```python
# Dane do zadania 7
x = sp.stats.uniform.rvs(loc=-3, scale=12, size=50, random_state=18)
y = -1.25*x + 7 + sp.stats.norm.rvs(loc=0, scale=3, size=50, random_state=67)
data = pd.DataFrame(data=np.array([x, y]).T, columns=["x", "y"])
data.head()
```

Output:
```
x          y
0  4.804491   4.379953
1  3.065440   3.600776
2  7.543218  -5.774798
3 -0.817917  11.715536
4  7.226797  -2.497470
```

### Zadanie 7: Estymacja przedziaowa funkcji regresji liniowej [12 pkt]

#### a ) Przygotowanie funkcji
Przygotuj funkcj `linear_regression_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziau ufnoci i przedziau predykcji regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego.
Zaprojektuj funkcj w taki spos贸b, aby obliczaa wartoci przedzia贸w dla 100 wartoci x r贸wno rozo偶onych od minimalnej do maksymalnej wartoci x z danych wejciowych.

<!-- grade_id: cell-04dbe3c917594f2b, points: 0, grade: false, solution: true -->
```python
def linear_regression_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "x" i "y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    Returns
    -------
    a: float
        Wsp贸czynnik r贸wnania regresji (slope).
    b: float
        Wsp贸czynnik r贸wnania regresji (intercept).
    results: DataFrame
        Tablica zawierajca domylny indeks oraz nastpujce kolumny:
        x - wartoci zmiennej niezale偶naj, dla kt贸rych okrelamy granice przedziau,
        y_reg - wartoci y wyliczone na podstawie wartoci x z wykorzystaniem wyliczonego r贸wnania regresji liniowej,
        ci_high - g贸rne granice przedziau ufnoci,
        ci_low - dolne granice przedziau ufnoci,
        pi_high - g贸rne granice przedziau predykcji,
        pi_low - dolne granice przedziau predykcji.
    """  
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ef82e28a9ad30a24, points: 7, grade: true, solution: false -->
```python
### Kom贸rka testowa
a, b , results = linear_regression_estimation(test_data_2, 0.05)
assert np.isclose(a, 2.072171435240703)
assert np.isclose(b, -4.6260980900060495)
assert results.shape == (100, 6)
assert len(set(list(results.columns)) & set(['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low'])) == 6
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].min().values, np.array([-1.8657939, -8.49234292, -7.06370958, -9.92097626, -4.39199716, -12.59268868]))) == True
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].max().values, np.array([7.4114666, 10.73173129, 12.12260347, 9.34085912, 14.81907365, 6.64438894]))) == True
assert np.all(np.isclose(results[['x', 'y_reg', 'ci_high', 'ci_low', 'pi_high', 'pi_low']].mean().values, np.array([2.77283635, 1.11969419, 2.13259929, 0.10678908, 5.09904685, -2.85965848]))) == True
```

#### b) Wykres 

Wykorzystaj przygotowan funkcj do wyznaczenia przedzia贸w na danych ze zmiennej `data` i przedstaw wyniki na pojedynczym wykresie. W obliczeniach przyjmij wybrany przez siebie poziom istotnoci. 

Przygotuj figur zgodnie z wytycznymi z nag贸wka notatnika. Dodatkowo:
- zawrzyj na wykresie punkty, na kt贸rych bya prowadzona analiza, prost regresji, granice przedziau ufnoci oraz granice przedziau predykcji,
- wyr贸偶nij poszczeg贸lne elementy wykresu (w tym znajdujce si na nim punkty) r贸偶nymi kolorami (granice poszczeg贸lnych przedzia贸w w tym samym kolorze) i formamy linii (ciga dla prostej regresjii, przerywane dla granic przedzia贸w),
- dodaj do figury legend opisujc wszystkie elementy wykresu. Umie w legendzie informacj o wyznaczonym r贸wnaniu regresji (wsp贸czynniki sformatowane do 4 miejsc po przecinku) oraz wartoci wsp贸czynnika ufnoci (sformatowane do 3 miejsc po przecinku) dla kt贸rego byy wyznaczane przedziay.

<!-- grade_id: cell-ecb421dfbf3d197e, points: 5, grade: true, solution: true -->
```python
# Wykres
# YOUR CODE HERE
raise NotImplementedError()
```

### Zadanie 8: Estymacja bootstapowa parametr贸w pr贸by losowej [7 pkt]

Przygotuj funkcj `bootstrap_parameter_estimation()`, kt贸ra korzystajc z metody bootstrap bdzie dokonywaa estymacji przedziaowej podstawowych parametr贸w pr贸by losowej - wartoci oczekiwanej, wariancji i odchylenia standardowego dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

W obliczeniach skorzystaj z nieobci偶onych estymator贸w wariancji i odchylenia standardowego.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`P nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie P (symbol lub nazwa estymowanego parametru), X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakadaj losowanie pr贸b wg nastpujcego schematu - kolejne pr贸by s losowane w pli for, a funkcja losujca pr贸b (np. `pd.sample()`) w kolejnych losowaniach ma ustawion warto parametru odpowiadajcego za ziarno generatora liczb losowych na warto `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwizaniu nie nale偶y korzysta z gotowych funkcji su偶cych do estymacji bootstrapowej.

<!-- grade_id: cell-93982a1f4c5f4ad1, points: 0, grade: false, solution: true -->
```python
def bootstrap_parameter_estimation(data, parameter, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej.
    parameter: str
        W zale偶noci od estymowanego parametru przyjmuje warto "mean", "var" lub "std".
    number_of_samples: int
        Liczba pr贸b bootstrapowych.
    alpha: float
        Warto okrelajca poziom istotnoci.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    parameter_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    parameter_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """       
    np.random.seed(random_state)
    n = len(data)
    estimates = np.zeros(number_of_samples)

    parameter_functions = {
        'mean': lambda sample: float(sample.mean()),
        'std': lambda sample: float(sample.std(ddof=1)),
        'var': lambda sample: float(sample.var(ddof=1))
    }

    for i in range(number_of_samples):
        sample = data.sample(n, replace=True, random_state=random_state + i)
        estimates[i] = parameter_functions[parameter](sample)

    sorted_estimates = np.sort(estimates)
    
    parameter_low = np.percentile(estimates, 100 * alpha / 2)
    parameter_high = np.percentile(estimates, 100 * (1 - alpha / 2))
    confidence_interval = 1 - alpha
    
    if parameter == 'mean':
        print(f"Mean (渭) nale偶y do przedziau [{parameter_low:.4f}, {parameter_high:.4f}] przy zao偶eniu poziomu ufnoci 1-伪 = {confidence_interval}")
    elif parameter == 'var':
        print(f"Variance (^2) nale偶y do przedziau [{parameter_low:.4f}, {parameter_high:.4f}] przy zao偶eniu poziomu ufnoci 1-伪 = {confidence_interval}")
    elif parameter == 'std':
        print(f"Standard Deviation () nale偶y do przedziau [{parameter_low:.4f}, {parameter_high:.4f}] przy zao偶eniu poziomu ufnoci 1-伪 = {confidence_interval}")
  
    return parameter_low, parameter_high
```

<!-- grade_id: cell-8d11ac5738ce3f52, points: 7, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(bootstrap_parameter_estimation(test_data_1, "mean", 1000, 0.05, 34), (4.913578501079754, 5.058598650238277)))
assert np.all(np.isclose(bootstrap_parameter_estimation(test_data_1, "var", 2000, 0.02, 31), (0.009780029754528753, 0.05012098474121563)))
assert np.all(np.isclose(bootstrap_parameter_estimation(test_data_1, "std", 1500, 0.1, 39), (0.11848597559111698, 0.20690074715252046)))
```

Output:
```
Mean (渭) nale偶y do przedziau [4.9136, 5.0586] przy zao偶eniu poziomu ufnoci 1-伪 = 0.95
Variance (^2) nale偶y do przedziau [0.0098, 0.0501] przy zao偶eniu poziomu ufnoci 1-伪 = 0.98
Standard Deviation () nale偶y do przedziau [0.1185, 0.2069] przy zao偶eniu poziomu ufnoci 1-伪 = 0.9
```

### Zadanie 9: Estymacja bootstrapowa za pomoc dedykowanych funkcji [6 pkt]

Korzystajc z funkcji `scipy.stats.bootstrap()` dokonaj estymacji przedziaowej nastpujcych parametr贸w na poziomie istotnoci 1-alpha = 0.98:
 - na podstawie zmiennej `test_data_1`: rednia, wariancja, odchylenie standardowe,
 - na podstawie zmiannej `test_data_2`: wsp贸czynnik korelacji liniowej Pearsona, wsp贸czynniki r贸wnania regresji liniowej a i b.
 
Wyniki zapisz do zmiennych o nazwach `mean_ci`, `var_ci`, `std_ci`, `r_ci`, `a_ci` oraz `b_ci` zawierajcych listy z dolnymi i g贸rnymi granicami przedzia贸w ufnoci. Podczas obliczania wartoci ka偶dego z parametr贸w ustaw warto argumentu funkcji `scipy.stats.bootstrap()` odpowiadajcego za ziarno generatora liczb losowych na 25.

Uwagi do wykonania zadania:
 - Automatyczne testy przewiduj obliczenie wartoci wariancji i odchylenia standardowego dla wartoci parametru `ddof=0`.

<!-- grade_id: cell-2b7dd70bd1282c22, points: 0, grade: false, solution: true -->
```python
# mean
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[49], line 2
      1 # mean
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-dc1bee32fb99d590, points: 0, grade: false, solution: true -->
```python
# var
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e1ec0872b87a59a, points: 0, grade: false, solution: true -->
```python
# std
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2b448c733aa676cc, points: 0, grade: false, solution: true -->
```python
# r
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e968f3a0c54deed1, points: 0, grade: false, solution: true -->
```python
# a
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0b198ef3de71c0e4, points: 0, grade: false, solution: true -->
```python
# b
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a2d98f9e8754891c, points: 6, grade: true, solution: false -->
```python
# Kom贸rka testowa
assert np.all(np.isclose(mean_ci, (4.890098551749027, 5.063500727815335)))
assert np.all(np.isclose(var_ci, (0.013001391902892368, 0.05670851692597062)))
assert np.all(np.isclose(std_ci, (0.11423521149909441, 0.23829111055240196)))
assert np.all(np.isclose(r_ci, (0.9156715777126913, 0.9822743695601276)))
assert np.all(np.isclose(a_ci, (1.792948341896073, 2.3282111601384945)))
assert np.all(np.isclose(b_ci, (-5.700934784848322, -3.4292283266002834)))
```