# Analiza danych przestrzennych - wiczenia laboratoryjne 2023/2024

Ten notatnik zalicza si do grupy zestaw贸w zada, na podstawie kt贸rych odbywa si zaliczenie wicze i podlega zwrotowi do oceny w ustalonym na zajciach terminie.

Uwagi i wytyczne og贸lne dotyczce uzupeniania i oceny notatnika:
- Podczas wykonywania zada nale偶y korzysta wycznie z pakiet贸w zaimportowanych na pocztku notatnika oraz z pakiet贸w wchodzcych w skad standardowej biblioteki Pythona, kt贸re mo偶na zaimportowa samodzielnie we wskazanej kom贸rce notatnika.
- Swoje rozwizania nale偶y wprowadza wycznie w miejce nastpujcych fragment贸w kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w kom贸rkach oznaczonych hasem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie nale偶y w 偶aden spos贸b modyfikowa pozostaych fragment贸w kodu oraz innych element贸w notatnika, w szczeg贸lnoci dodawa lub usuwa kom贸rek oraz zmienia nazwy pliku.
- Je偶eli zestaw zada wymaga skorzystania z fragment贸w kodu opracowanego w ramach wczeniejszych zestaw贸w zada nale偶y je umieci we wskazanej kom贸rce notatnika.
- Otrzymywane wyniki i odpowiedzi maj by rezultatem wykonania napisanego kodu, odpowiedzi uzupeniane manualnie nie podlegaj ocenie.
- Zawarte w notatniku automatyczne testy maj charakter pogldowy. Dotycz one wybranych aspekt贸w zada i maj na celu wyapa podstawowe bdy. Przejcie przez kod wszystkich test贸w nie oznacza, 偶e zadanie jest wykonane w caoci poprawnie i zostanie ocenione na maksymaln liczb punkt贸w.
- Zadania nale偶y wykona w taki spos贸b, aby podczas wykonywania kodu nie zostay wywietlone 偶adne ostrze偶enia.
- Zadania, kt贸re powoduj wywietlenie komunikatu o bdzie przerywajcym wykonywanie kodu nie podlegaj ocenie.

Uwagi i wytyczne og贸lne dotyczce wizualizacji wynik贸w:
- Wszystkie wykresy powinny by wykonane w jednolitym, przejrzystym i czytelnym stylu, posiada odpowiednio dobrane proporcje i zakresy wartoci osi.
- Figury powinny mie ustawione biae to, tak, aby niezale偶nie od ustawie notatnika wszystkie elementy wykres贸w byy dobrze widoczne (domylnie to jest prze藕roczyste co mo偶e powodowa problemy w notatnikach z ustawionym ciemnym tem). Rozmiar poziomy figur nie powinien przekracza 20 cali.
- Wykresy oraz ich osie powinny mie nadane tytuy. Je偶eli w obrbie figury znajduje si wicej ni偶 jeden wykres to figura r贸wnie偶 powinna mie nadany tytu.

Przed odesaniem zestawu zada do oceny prosz uzupeni kom贸rk z danymi autor贸w rozwizania (nazwa zespou oraz imiona, nazwiska i numery indeks贸w czonk贸w zespou) oraz upewni si, 偶e notatnik dziaa zgodnie z oczekiwaniami. W tym celu nale偶y skorzysta z opcji **Restart Kernel and Run All Cells...** dostpnej na g贸rnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespou: Zuzanna Brauer, Sandra Gowacka, Zuzanna Chmielarska <br>
Czonkowie: IAD13

---

## Zestaw zada 4: Wybrane teoretyczne rozkady prawdopodobiestwa oraz elementy teorii estymacji (Notatnik 2/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
from math import sqrt
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[32], line 2
      1 from math import sqrt
----> 2 from sklearn.linear_model import LinearRegression

ModuleNotFoundError: No module named 'sklearn'
```

### Dane do zada

W celu wygenerowania danych wykorzystywanych w zawartych w notatniku zadaniach i kom贸rkach testowych wykonaj obie poni偶sze kom贸rki.

```python
# Dane do test贸w 1
test_data_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
test_data_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
# Dane do test贸w 2
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=250, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=250, random_state=13)
test_data_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["X", "Y"])
test_data_2.head()
```

Output:
```
X         Y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

### Zadanie 3: Estymacja przedziaowa wartoci oczekiwanej [7 pkt]

Przygotuj funkcj `mean_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci oczekiwanej dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

` nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-3162b1779d0c603e, points: 0, grade: false, solution: true -->
```python
def mean_estimation(data, alpha, population_std="unknown"):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej,
    alpha: float
        Warto okrelajca poziom istotnoci.
    population_std: str or float
        Je偶eli odchylenie standardowe populacji nie jest znane to parametr przyjmuje warto "unknown",
        w przeciwnym wypadku jest to warto okrelajca odchylenie standardowe populacji.
    
    Returns
    -------
    mean_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    mean_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
   
    n = len(data)
    data = data.to_numpy() # zamiana formatu poniewa偶 problem z printem przy pd.Series
    mean = np.mean(data)

    if population_std == "unknown":
        t_score = sp.stats.t.ppf(1 - alpha/2, df = n-1)
        margin = t_score*np.std(data, ddof = 1)/np.sqrt(n)
    else:
        t_score = sp.stats.norm.ppf(1 - alpha/2)
        margin = t_score*population_std /np.sqrt(n)
        
    mean_low = mean - margin
    mean_high = mean + margin
    print(f" nale偶y do przedziau [{mean_low:.4f}, {mean_high:.4f}] przy zao偶eniu poziomu ufnoci 1 - {alpha} = {1 - alpha:.2f}")

    return (mean_low,mean_high)
```

<!-- grade_id: cell-37f4b7bd74d39411, points: 7, grade: true, solution: false -->
```python
### Kom贸rka testowa

# Test na test_data_1 przy zao偶eniu, 偶e znane jest odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.1, population_std=0.2), (4.905619719495418, 5.05273990041144))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.05, population_std=0.2), (4.891527555895771, 5.066832064011087))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.02, population_std=0.2), (4.875142370240081, 5.083217249666777))))
print()

# # Test na test_data_1 przy zao偶eniu, 偶e nie jest znane odchylenie standardowe populacji, z kt贸rej zostaa pobrana pr贸bka losowa
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.1), (4.913356369231727, 5.045003250675131))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.05), (4.899504007229391, 5.058855612677467))))
print()
assert(np.all(np.isclose(mean_estimation(test_data_1, 0.02), (4.882508507351398, 5.07585111255546))))
```

Output:
```
nale偶y do przedziau [4.9056, 5.0527] przy zao偶eniu poziomu ufnoci 1 - 0.1 = 0.90

 nale偶y do przedziau [4.8915, 5.0668] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95

 nale偶y do przedziau [4.8751, 5.0832] przy zao偶eniu poziomu ufnoci 1 - 0.02 = 0.98

 nale偶y do przedziau [4.9134, 5.0450] przy zao偶eniu poziomu ufnoci 1 - 0.1 = 0.90

 nale偶y do przedziau [4.8995, 5.0589] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95

 nale偶y do przedziau [4.8825, 5.0759] przy zao偶eniu poziomu ufnoci 1 - 0.02 = 0.98
```

### Zadanie 4: Estymacja przedziaowa wariancji [5 pkt]
Przygotuj funkcj `variance_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wariancji dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`^2  nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-1f107aa4297b1e13, points: 0, grade: false, solution: true -->
```python
def variance_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i pojedyncz kolumn "X" z wartociami pr贸by losowej,
    alpha: float
        Warto okrelajca poziom istotnoci.
    
    Returns
    -------
    var_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    var_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """
    
    x = data.to_numpy()
    n = len(x)
    mean = np.mean(data)
    var = np.sum((x - mean)*(x - mean))/n
    var_low = var*n/sp.stats.chi2.ppf(1 - alpha/2, n - 1)
    var_high = var*n/sp.stats.chi2.ppf(alpha/2, n - 1)
    print(f"^2 nale偶y do przedziau [{var_low:.4f}, {var_high:.4f}] przy zao偶eniu poziomu ufnoci 1 - {alpha} = {1 - alpha:.2f}")
    
    return (var_low,var_high)
```

<!-- grade_id: cell-bd878d13cb44f436, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.1), (0.018268125369359626, 0.05442967510189219))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.05), (0.0167618487538223, 0.06182728513230527))));
print()
assert(np.all(np.isclose(variance_estimation(test_data_1, 0.02), (0.015215598500840796, 0.07214532146342346))));
```

Output:
```
^2 nale偶y do przedziau [0.0183, 0.0544] przy zao偶eniu poziomu ufnoci 1 - 0.1 = 0.90

^2 nale偶y do przedziau [0.0168, 0.0618] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95

^2 nale偶y do przedziau [0.0152, 0.0721] przy zao偶eniu poziomu ufnoci 1 - 0.02 = 0.98
```

### Zadanie 5: Estymacja wartoci wsp贸czynnika korelacji liniowej [5 pkt]

Przygotuj funkcj `correlation_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnika korelacji liniowej Pearsona dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`r nale偶y do przedziau [X, Y] przy zao偶eniu poziomu ufnoci 1- = Z`

gdzie X, Y i Z s automatycznie uzupeniane przez funkcj, a X i Y dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-4eef8786fb47018a, points: 0, grade: false, solution: true -->
```python
def correlation_estimation(data, alpha):

    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziau ufnoci.
    r_corr_high: float
        G贸rna granica wyliczonego przedziau ufnoci.
    """   

    r = sp.stats.pearsonr(data['X'], data['Y']).statistic
    r_corr_low = r - sp.stats.norm.ppf(1 - alpha/2)*(1 - r*r)/sqrt(len(data))
    r_corr_high = r + sp.stats.norm.ppf(1 - alpha/2)*(1 - r*r)/sqrt(len(data))
    print(f"r nale偶y do przedziau [{r_corr_low:.4f}, {r_corr_high:.4f}] przy zao偶eniu poziomu ufnoci 1 - {alpha} = {1 - alpha:.2f}")

    return(r_corr_low, r_corr_high)
```

<!-- grade_id: cell-361fdd010375f603, points: 5, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(correlation_estimation(test_data_2, 0.05), (0.9494696249316307, 0.9692185148958979)))
assert np.all(np.isclose(correlation_estimation(test_data_2, 0.02), (0.9476237555239514, 0.9710643843035772)))
```

Output:
```
r nale偶y do przedziau [0.9495, 0.9692] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95
r nale偶y do przedziau [0.9476, 0.9711] przy zao偶eniu poziomu ufnoci 1 - 0.02 = 0.98
```

### Zadanie 6: Estymacja wsp贸czynnik贸w r贸wnania regresji liniowej [8 pkt]

Przygotuj funkcj `interval_linear_regression_coefficients_estimation()`, kt贸ra bdzie dokonywaa estymacji przedziaowej wartoci wsp贸czynnik贸w r贸wnania regresji liniowej dla danych wejciowych w postaci szeregu szczeg贸owego zgodnie ze schematem z zaczonego do notatnika zestawu wzor贸w.

Opr贸cz zwracanych wartoci granic przedziau funkcja powinna wywietla nastpujcy komunikat:

`a nale偶y do przedziau [X1, Y1] przy zao偶eniu poziomu ufnoci 1- = Z1` </br>
`b nale偶y do przedziau [X2, Y2] przy zao偶eniu poziomu ufnoci 1- = Z2`

gdzie X1, X2, Y1, Y2, Z1 i Z2 s automatycznie uzupeniane przez funkcj, a X1, X2, Y1 i Y2 dodatkowo sformatowane w taki spos贸b, 偶eby wywietlay si z dokadnoci do 4 miejsc po przecinku.

<!-- grade_id: cell-853c41ec96cf7c37, points: 0, grade: false, solution: true -->
```python
def linear_regression_coefficients_estimation(data, alpha):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierajca domylny indeks i dwie kolumny "X" i "Y" z wynikami pr贸by losowej.
    alpha: float
        Warto okrelajca poziom istotnoci.
    
    Returns
    -------
    a_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru a (a_low, a_high).
    b_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziau ufnoci parametru b (b_low, b_high).
    """   

    reg = sp.stats.linregress(data['X'], data['Y'])
    a_low = reg.slope - sp.stats.t.ppf(1 - alpha/2, df = len(data) - 2)*reg.stderr
    a_high = reg.slope + sp.stats.t.ppf(1 - alpha/2, df = len(data) - 2)*reg.stderr
    b_low = reg.intercept - sp.stats.t.ppf(1 - alpha/2, df = len(data) - 2)*reg.intercept_stderr
    b_high = reg.intercept + sp.stats.t.ppf(1 - alpha/2, df = len(data) - 2)*reg.intercept_stderr
    a_ci = (a_low, a_high)
    b_ci = (b_low, b_high)
    print(f"a nale偶y do przedziau [{a_low:.4f}, {a_high:.4f}] przy zao偶eniu poziomu ufnoci 1 - {alpha} = {1 - alpha:.2f}")
    print(f"b nale偶y do przedziau [{b_low:.4f}, {b_high:.4f}] przy zao偶eniu poziomu ufnoci 1 - {alpha} = {1 - alpha:.2f}")
    
    return(a_ci, b_ci)
```

<!-- grade_id: cell-3d47e4e5f2b4d455, points: 8, grade: true, solution: false -->
```python
### Kom贸rka testowa
assert np.all(np.isclose(linear_regression_coefficients_estimation(test_data_2, 0.1), ((1.90330932143819, 2.024457061652292), (-5.331039505544365, -4.7958291264947706))))
print()
assert np.all(np.isclose(linear_regression_coefficients_estimation(test_data_2, 0.05), ((1.8916219551002786, 2.0361444279902035), (-5.382672328053067, -4.744196303986069))))
```

Output:
```
a nale偶y do przedziau [1.9033, 2.0245] przy zao偶eniu poziomu ufnoci 1 - 0.1 = 0.90
b nale偶y do przedziau [-5.3310, -4.7958] przy zao偶eniu poziomu ufnoci 1 - 0.1 = 0.90

a nale偶y do przedziau [1.8916, 2.0361] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95
b nale偶y do przedziau [-5.3827, -4.7442] przy zao偶eniu poziomu ufnoci 1 - 0.05 = 0.95
```

```python

```